<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/Test (dataprog)</text>
    </category>
    <info format="html">
      <text><![CDATA[<p>Question for testing</p>]]></text>
    </info>
    <idnumber></idnumber>
  </question>

<!-- question: 208  -->
  <question type="coderunner">
    <name>
      <text>Test av Coderunner</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">Print ut "Hei coderunner"<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[print("Hei coderunner")]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>### Sjekk output</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>Hei coderunner
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 210  -->
  <question type="truefalse">
    <name>
      <text>LTI spm</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">Fungerer LTI verktøyet?<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>1.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <answer fraction="0" format="moodle_auto_format">
      <text>true</text>
      <feedback format="html">
        <text><![CDATA[<p dir="ltr" style="text-align: left;">FLOTT!<br></p>]]></text>
      </feedback>
    </answer>
    <answer fraction="100" format="moodle_auto_format">
      <text>false</text>
      <feedback format="html">
        <text><![CDATA[<p dir="ltr" style="text-align: left;">NEEEEI!<br></p>]]></text>
      </feedback>
    </answer>
  </question>

<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/Introduksjon-variabler-regneartene-funksjoner</text>
    </category>
    <info format="html">
      <text>Spørsmål om variabler, enkle datatyper og regneartene</text>
    </info>
    <idnumber></idnumber>
  </question>

<!-- question: 311  -->
  <question type="coderunner">
    <name>
      <text>formater-hundretusen</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[Skriv et program som ber om et tall over 100 tusen via `input`-funksjonen,<br>
og printer ut tallet som et heltall rundet av til nærmeste 10 tusen]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Vi henter inn et tall med `input`-funksjonen - det blir lagret som en tekststreng<br>
Deretter gjør vi det om til et flyttall med `float()`, runder det av til nærmeste <br>
10 tusen ved å bruke `round(tall, -4)`.  Når vi gir `round` negative verdier for "desimalplassen" <br>
blir tallet rundet av til plassen *foran* komma. 4 siffer foran komma er 10 tusen plassen. <br>
Til slutt bruker vi `int` til å gjøre tallet om til et heltall, lagrer det hele i en variabel og printer ut en liten melding <br>

```python
{{QUESTION.answer}}
```]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate>0</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[{"answer_code":["tall_str = input(\"Gi meg et tall over 100 tusen\\n\")\ntall_float = float(tall_str)\ntall_avrundet = int( round(tall_float,-4))\nprint(\"Du gav meg ca\", tall_avrundet)"],"test_code":[""],"show_hide":[""],"prefix_ans":["1"]}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>RegexGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>1</twigall>
    <uiplugin>scratchpad</uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 1</text>
      </testcode>
      <stdin>
                <text>123576</text>
      </stdin>
      <expected>
                <text>120000[^\.]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 2</text>
      </testcode>
      <stdin>
                <text>126487</text>
      </stdin>
      <expected>
                <text>130000[^\.]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 3</text>
      </testcode>
      <stdin>
                <text>4666000</text>
      </stdin>
      <expected>
                <text>4670000[^\.]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 305  -->
  <question type="coderunner">
    <name>
      <text>Input-dobling</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[Skriv et program som spør bruker etter et heltall med `input`-funksjonen, <br>
dobler heltallet og printer det ut]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Vi bruker `input` funksjonen til å få en verdi fra brukeren<br>
```python
variabel = input("tekst som skal vises til bruker")
```
Inputfunksjonen gir alltid tilbake en tekststreng. <br>
Vi kan gjøre den om til et heltall med `int()`,<br>
men dersom tekstrengen ikke er et gyldig heltall, vil programmet kræsje<br>

Når vi har hentet inn heltallet kan vi gange det med 2 og printe ut]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[{{ STUDENT_ANSWER }}

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

    
def input(x):
    print("ble en kjørt?")
    return "{{Test.stdin}}"

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate>0</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[tall_str = input("Gi meg et heltall, så skal jeg doble det\n")
tall_int = int(tall_str)
print("Det dobbelte er:", tall_int*2)]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>RegexGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 1</text>
      </testcode>
      <stdin>
                <text>5
</text>
      </stdin>
      <expected>
                <text>10[^\.]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 2</text>
      </testcode>
      <stdin>
                <text>7</text>
      </stdin>
      <expected>
                <text>14[^\.]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 3</text>
      </testcode>
      <stdin>
                <text>125</text>
      </stdin>
      <expected>
                <text>250[^\.]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 300  -->
  <question type="coderunner">
    <name>
      <text>Input-meter-til-km</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[Skriv et program som spør bruker etter en lengde i meter med `input` funksjonen <br>
og printer ut lengden i kilometer med 1 desimal etter komma. <br>
Bruk `round` funksjonen til å runde av til riktig antall desimaler]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Vi bruker `input` funksjonen til å få en verdi fra brukeren<br>
```python
variabel = input("tekst som skal vises til bruker")
```
Inputfunksjonen gir alltid tilbake en tekststreng. <br>
Vi kan gjøre den om til en numerisk verdi med `int()` eller `float()`,<br>
men dersom tekstrengen er et desimaltall feks "3.14", *må* vi bruke `float()`<br>
Vi gjør deretter meter om til kilometer ved å dele på 1000, og avrunder med round:
```python
kilometer_avrundet = round(kilometer,1) # Runder av til 1 tall bak komma
```]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[{{ STUDENT_ANSWER }}

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

called_round = False
round_arg = None
round_tmp = round
def round(x,des):
    called_round = True
    round_arg = des
    return round_tmp(x,des)
    
def input(x):
    return {{Test.stdin}}

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
called_round = False
round_arg = None
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate>0</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[meter = input("Gi meg en lengde i meter\n")
meter_float = float(meter)
kilometer = meter_float/1000
kilometer_avrundet = round(kilometer,1)
print(meter, " meter er ", kilometer_avrundet, "kilometer")]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>RegexGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 1</text>
      </testcode>
      <stdin>
                <text>12400
</text>
      </stdin>
      <expected>
                <text>12\.4[^0-9]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 2</text>
      </testcode>
      <stdin>
                <text>12440</text>
      </stdin>
      <expected>
                <text>12\.4[^0-9]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 3</text>
      </testcode>
      <stdin>
                <text>12550</text>
      </stdin>
      <expected>
                <text>12\.6[^0-9]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Test 4</text>
      </testcode>
      <stdin>
                <text>650</text>
      </stdin>
      <expected>
                <text>0\.7[^0-9]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 1065  -->
  <question type="coderunner">
    <name>
      <text>Regneartene</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[Skriv et program som tar to tall fra bruker og printer ut:
<ul>
   <li> Summen av tallene (plusse) </li>
   <li> Differansen (trekk fra hverandre) </li>
   <li> Produktet (ganget med hverandre) </li>
   <li> Og kvotienten (første tall dele på andre tall) </li>
</ul>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text>Vi tar først 2 tall fra bruker med `input`-funksjonen, og husker på å gjøre de til flyttall med `float()`
Resten burde være selvforklarende:
```python
{{QUESTION.answer}}
```</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[{{ STUDENT_ANSWER }}

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate>0</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[tall1 = float( input("Gi meg et tall\n") )
tall2 = float( input("Gi meg et tall til\n") )

summe = tall1+tall2
diff = tall1-tall2
produkt = tall1*tall2
kvotient = tall1/tall2

print(tall1, "+", tall2, "=", summe)
print(tall1, "-", tall2, "=", diff)
print(tall1, "ganger", tall2, "=", produkt)
print(tall1, "dele op", tall2, "=", kvotient)]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>RegexGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>1</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Addisjon</text>
      </testcode>
      <stdin>
                <text>12
14</text>
      </stdin>
      <expected>
                <text>26</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Differanse</text>
      </testcode>
      <stdin>
                <text>30
23</text>
      </stdin>
      <expected>
                <text>7</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Produkt</text>
      </testcode>
      <stdin>
                <text>8
8</text>
      </stdin>
      <expected>
                <text>64</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>#Kvotient</text>
      </testcode>
      <stdin>
                <text>9
3</text>
      </stdin>
      <expected>
                <text>3</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 255  -->
  <question type="coderunner">
    <name>
      <text>variabler-one-liner</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Skriv et program som lager 2 variabler <code> a </code> og <code> b </code>, med verdier 1 og <code> True </code>. <br>
Programmet kan bare være 1 linje langt!
</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text>Vi kan lage flere variabler på 1 linje ved å separere variabelnavn og verdier med et komma (,)
```python
var1, var2, var3 = verdi1, verdi2, verdi3
```</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[{{ STUDENT_ANSWER }}

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

antall_linjer = __student_answer__.count('\n')+1
if __student_answer__.splitlines()[-1] == '\n':
    antall_linjer -= 1

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>a,b = 1, True</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(a)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>1</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(b)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>True</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text><![CDATA["Antall linjer i programmet ditt"
print(antall_linjer)]]></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>1
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 228  -->
  <question type="coderunner">
    <name>
      <text>variabler-print-variabel</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Programmet under skal printe ut verdien av variabel <code> a </code>, men her er gjort en feil.
Fiks programmet.]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<link rel="stylesheet"
        href="https://cdnjs.cloudflare.com/ajax/libs/highlight.js/10.0.3/styles/default.min.css">
  <script src="https://cdnjs.cloudflare.com/ajax/libs/highlight.js/10.0.3/highlight.min.js"></script>
  <script>hljs.initHighlightingOnLoad();</script>
<p>
Problemet med programmet er at <code class="python"> print(a) </code> kommandoen skjer før variabelen <code class="python"> a </code> lages. <br>
Vi fikser problemet ved å kalle <code> print(a) </code> <it> etter </it> at variabelen vi printer ut lages
</p>
<pre><code class="python">
a = 1
b= "Rubber baby bubby bunkers"
c = 1e-3
print(a)
</code></pre>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[print(a)

a = 1
b= "Rubber baby bubby bunkers"
c = 1e-3]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[a = 1
b= "Rubber baby bubby bunkers"
c = 1e-3
print(a)
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>2</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text><![CDATA["Kjører og sjekker"]]></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>1
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 1061  -->
  <question type="coderunner">
    <name>
      <text>variabler-type-boolean</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Skriv et lite program som lager en variabel med navn <code> var </code> som inneholder en boolsk verdi]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Vi lager variabler med `variabelnavn = <verdi>` <br>
Hvilken type variablen er bestemmes av verdien vi gir variabelen.<br>
Skal den ha typen `bool` -  - kan vi gi verdien med:<br>
```python
var1 = True
var2 = False
var3 = bool(0.0) # False
var4 = bool(1) # True
```
Funksjonen `bool()` tar en numerisk verdi og gir `False` for 0, eller 0.00 - og `True` ellers]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[tmp = print
stud_logs = []

def hijack_print(*args, **kwargs):
   to_print = [a for a in args]
   stud_logs.append(to_print)

print = hijack_print

{{ STUDENT_ANSWER }}

print=tmp

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>var = True</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(type(var))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text><![CDATA[<class 'bool'>
]]></text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 1059  -->
  <question type="coderunner">
    <name>
      <text>variabler-type-float-funk</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Skriv et lite program som lager en variabel med navn <code> var </code> som inneholder et flyttall,
men uten bruk av et eneste <code> . </code> (punktum) i koden]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Vi lager variabler med `variabelnavn = <verdi>` <br>
Hvilken type variablen er bestemmes av verdien vi gir variabelen.<br>
Skal den ha typen `float` - et flyttall, kunne vi gitt den et desimaltall, feks 3.14 <br>
Uten å bruke punktum kan vi bruke `float()` funksjonen eller vitenskaplig notasjon
```python
var1 = float(4) # Gir et flyttall 4.0000....
var2 = 3e6  # 3 millioner som flyttall 
```]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[tmp = print
stud_logs = []

def hijack_print(*args, **kwargs):
   to_print = [a for a in args]
   stud_logs.append(to_print)

print = hijack_print

{{ STUDENT_ANSWER }}

print=tmp

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>var = float(1)</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(type(var))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text><![CDATA[<class 'float'>
]]></text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text><![CDATA["Inneholder koden '.'?"
print(__student_answer__.count('.'))]]></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>0</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 1057  -->
  <question type="coderunner">
    <name>
      <text>variabler-type-float-std</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Skriv et lite program som lager en variabel med navn <code> var </code> som inneholder et flyttall]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Vi lager variabler med `variabelnavn = <verdi>` <br>
Hvilken type variablen er bestemmes av verdien vi gir variabelen.<br>
Skal den ha typen `float` - et flyttall - kan vi gi verdien med:<br>
```python
var1 = 1.0 # Vi gir variablen et desimaltall
var2 = float(42) # funksjonen float() gjør om tekst eller heltall til flyttall
var3 = 4e6 # Vitenskaplig notasjon er også flyttall
```]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[tmp = print
stud_logs = []

def hijack_print(*args, **kwargs):
   to_print = [a for a in args]
   stud_logs.append(to_print)

print = hijack_print

{{ STUDENT_ANSWER }}

print=tmp

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>var = 1.0</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(type(var))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text><![CDATA[<class 'float'>
]]></text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 1055  -->
  <question type="coderunner">
    <name>
      <text>variabler-type-int</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Skriv et lite program som lager en variabel med navn <code> var </code> som inneholder et heltall]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Vi lager variabler med `variabelnavn = <verdi>` <br>
Hvilken type variablen er bestemmes av verdien vi gir variabelen.<br>
Skal den ha typen `int` - et heltall - kan vi gi verdien med:<br>
```python
var1 = 1 #Vi oppgir et tall uten desimaler
var2 = int("1") # funksjonen int() gjør om tekst til tall (om det er et gyldig heltall)
var3 = int(4.37) # Gir vi int()-funksjonen et flyttall kutter den desimalene (ingen avrunding)
```]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[tmp = print
stud_logs = []

def hijack_print(*args, **kwargs):
   to_print = [a for a in args]
   stud_logs.append(to_print)

print = hijack_print

{{ STUDENT_ANSWER }}

print=tmp

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>var = 1</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(type(var))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text><![CDATA[<class 'int'>
]]></text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 1053  -->
  <question type="coderunner">
    <name>
      <text>variabler-type-str</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Skriv et lite program som lager en variabel med navn <code> var </code> som inneholder en tekststreng]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Vi lager variabler med `variabelnavn = <verdi>` <br>
Hvilken type variablen er bestemmes av verdien vi gir variabelen.<br>
Skal den ha typen `str` - en tekststreng - kan vi gi verdien med:<br>
```python
var1 = "din tekststreng"
var2 = 'din tekststreng'
var3 = """Din tekststreng,
som går over
mange linjer"""
var4 = '''Tekststreng
over 
flere linjer
'''
var5 = str(4556) # gir tekststrengen "4556"
```]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[tmp = print
stud_log = []

def hijack_print(*args, **kwargs):
   to_print = [a for a in args]
   stud_log.append(to_print)


print = hijack_print

{{ STUDENT_ANSWER }}

print=tmp
__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[var = "Fire franske fruer fiser full fart forover"]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(type(var))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text><![CDATA[<class 'str'>
]]></text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/Intro-modules-loops-lists</text>
    </category>
    <info format="html">
      <text></text>
    </info>
    <idnumber></idnumber>
  </question>

<!-- question: 355  -->
  <question type="coderunner">
    <name>
      <text>For-løkke: Doble annethvert element i liste</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text>Doble annethvert element i listen `a` ved å bruke en `for`-løkke og `range`.
Det første elementet skal altså dobles, det andre skal forbli det samme, det tredje skal dobles igjen, osv...</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text>* For å utføre oppgaven bruker vi itererer vi over indeksområdet til `a` med en `for`-løkke og `range`-funksjonen
* Vi trenger å vite lengden til listen, den finner vi med: `n = len(a)`
* I løkken er bruker vi da `range(0,n,2)` som gir oss tall fra og med 0 til n med steglengde 2: 0,2,4,....,n-1

For-løkken blir da slik:
```python
for i in range(0,n,2):
   a[i] = 2*a[i]
``</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload>a = [1,2,3,4,5,56]</answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[import re

{{ STUDENT_ANSWER }}

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

exec_string = re.search(r'a *= *(\[.*\])', __student_answer__)
new_ans = re.sub(exec_string.group(1).replace("[", "\\[").replace("]", "\\]"), "[5,3,7,1,7,9,45,89,3,23423,45]", __student_answer__)

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>a = [1,2,3,4,5,56]
n = len(a)
for i in range(0,n,2):
   a[i] = 2*a[i]</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(a)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[2, 2, 6, 4, 10, 56]
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>exec(new_ans)
print(a)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[10, 3, 14, 1, 14, 9, 90, 89, 6, 23423, 90]
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 340  -->
  <question type="coderunner">
    <name>
      <text>For-loop intro</text>
    </name>
    <questiontext format="markdown">
      <text>Print ut elementene i liste `a` ved bruk av en `for`-løkke</text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[a = ["hei", 1,2, 3.54e5]]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[a = ["hei", 1,2, 3.54e5]
for element in a:
    print(element)]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text># Sjekker elementer</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>hei
1
2
354000.0
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text><![CDATA[print(__student_answer__.count("print") < 2)]]></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>True
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 328  -->
  <question type="coderunner">
    <name>
      <text>Function: average</text>
    </name>
    <questiontext format="markdown">
      <text>Sjekk ut funksjonene `len(«liste»)` og `sum(«liste»)` og bruk disse til å lage en funksjon `gjennomsnitt` som tar en liste og returnerer gjennomsnittsverdien til listen</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text>*  `len(«liste»)` gir lengden, eller antall elementer til en liste
*  `sum(«liste»)` git summen av alle elementene i listen
*  Vi kan bruke disse til å lage en funksjon som regner ut gjennomsnittet til en liste slik:
```python3
def gjennomsnitt(a):
   n = len(a) #Antall elementer
   a_sum = sum(a) #Summen av alle elementer
   return a_sum/n
```</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[sum_tmp = sum
sum_called = False
def sum_hijack(a):
    global sum_called
    sum_called = True
    return sum_tmp(a)

len_tmp = len
len_called = False
def len_hijack(a):
    global len_called
    len_called = True
    return len_tmp(a)

len = len_hijack
sum = sum_hijack
    

{{ STUDENT_ANSWER }}

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>def gjennomsnitt(a):
   n = len(a) #Antall elementer
   a_sum = sum(a) #Summen av alle elementer
   return a_sum/n</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(gjennomsnitt([1,2,3,4,5]))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>3.0
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(gjennomsnitt([7,6,5,4,34,3]))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>9.833333333333334
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(sum_called, len_called)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>True True
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 318  -->
  <question type="coderunner">
    <name>
      <text>Function: doble</text>
    </name>
    <questiontext format="markdown">
      <text>Lag en funksjon `doble` som tar et tall og returnerer det dobbelte av tallet</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text>Vi lager en egen funksjon slik:
```python
def funksjonsnavn(arg1, arg2):
   #Kode som gjør noe
   #Mer kode
   return resultat
```
* `def` betyr nå skal vi definere en funksjon
* `funksjonsnavn` er hvilket som helst gyldig navn til funksjonen
* `(arg1, arg2,...)` er listen med argumenter som funksjon skal ta når den kalles
* Vi husker `:` kolon etter argumentlisten
* Koden funksjonen skal kjøre trenger et innrykk
* Vi returnerer resultatet med `return`

```python3
def doble(x):
   dobbelt = 2*x
   return dobbelt
```</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[printt = print


stud_log =[]
def print_fake(*args):
    for arg in args:
        stud_log.append(arg)

print = print_fake

{{ STUDENT_ANSWER }}
__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
printt(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>def doble(tall):
    return tall*2</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>printt(doble(2))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>4
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>printt(doble(3))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>6
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>printt(doble(4))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>8
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 334  -->
  <question type="coderunner">
    <name>
      <text>Function: Function returning function</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text>Lag en funksjon `lagRenteFunksjon` som tar et tall \(r\) og returnerer en funksjon som tar et tall \(P\)og et tall \(t\) og returnerer \(P\cdot (1+r)^{t}\)</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text>* Vi kan definere funksjoner inne i funksjoner og returnere disse som om de var en vanlig variabel
* I dette tilfellet lager vi en funksjon som tar en rentesats `r` og returnerer en funksjon vi kan bruke til å regne ut sluttverdi på et beløp `P` etter `t` år forrrentning
```python
def lagRenteFunksjon(r):
    def renteFunksjon(P,t):
        return P*(1+r)**t
    return renteFunksjon
```</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload>def lagRenteFunksjon(r):
    # ??
    #
    return ???</answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>def lagRenteFunksjon(r):
    def renteFunksjon(P,t):
        return P*(1+r)**t
    return renteFunksjon</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>f = lagRenteFunksjon(0.1)
print(round(f(1000,10),1))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>2593.7
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>f = lagRenteFunksjon(0.05)
print(round(f(10000,5),1))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>12762.8
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 323  -->
  <question type="coderunner">
    <name>
      <text>Function: hello</text>
    </name>
    <questiontext format="markdown">
      <text>Lag en funksjon `hilse(navn)` som tar et navn og printer ut en hilsen som inneholder navnet.
Funksjonen skal kun printe ut 1 linje</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text><![CDATA[Funksjonen vår skal kun printe ut 1 tekststreng som inneholder navnet som ble gitt.
Funksjonen skal ikke returnere noen verdi -- det er det heller ikke krav om at en funksjon MÅ gjøre så vi kan bare droppe `return` i funksjonsblokken

```python
def hilse(navn):
   print("Hei på deg", navn)
```]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def hilse(navn):
   print("Være hilset kjære", navn)
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>RegexGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text><![CDATA[print(type(hilse("jonas")))]]></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>NoneType</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text><![CDATA[hilse("jonas")]]></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>jonas</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 338  -->
  <question type="coderunner">
    <name>
      <text>Function: Multiple return values</text>
    </name>
    <questiontext format="markdown">
      <text>Lag en funksjon  `func` som tar en liste og returnerer summen av listen **og** største verdi i listen</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text>* Vi kan bruke `sum` og `max` til å finne summen av- og største element i en listen
* For å returnere to verdier, si a og b, fra en funksjon skriver vi når vi vil returnere: `return a, b`</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[{{ STUDENT_ANSWER }}

def func1_t(a):
   sum_a = sum(a)
   max_a = max(a)
   return sum_a, max_a
  
def func2_t(a):
    sum_a = sum(a)
    max_a = max(a)
    return max_a, sum_a
    
liste1 = [1,2,3,4,5]
liste2 = [2,5,55,22,5,7,8,90,2]

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>def func(a):
   sum_a = sum(a)
   max_a = max(a)
   return max_a, sum_a</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>s1, s2 = func(liste1)
t1, t2 = func1_t(liste1)
korrekt = (s1,s2) == (t1,t2) or (s1,s2) == (t2,t1)
print(korrekt)
</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>True
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>s1, s2 = func(liste2)
t1, t2 = func1_t(liste2)
korrekt = (s1,s2) == (t1,t2) or (s1,s2) == (t2,t1)
print(korrekt)
</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>True</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 352  -->
  <question type="coderunner">
    <name>
      <text>Lister: 2 siste</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text>Fullfør funksjonen `to_siste` som tar en liste a og printer ut de to siste elementene.</text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[Vi kan indeksere lister "bakenifra" ved å bruke negative indekser. `liste[-1]` gir siste element, `liste[-2]` gir nest siste element, osv...]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload>def to_siste(a):
    print(?????)
    print(?????)</answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>def to_siste(a):
    print(a[-1])
    print(a[-2])</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>to_siste([1,2,3,4,5,6,7])</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>7
6
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>to_siste([11,5,7,32,76,34,34,28,29])</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>29
28
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 357  -->
  <question type="coderunner">
    <name>
      <text>Lister av lister loop</text>
    </name>
    <questiontext format="markdown">
      <text>* I programmet under har vi en `kundeliste` variabel som består av en liste fylt med kundeinfo
* Kundeinformasjonen er i seg selv en liste, og består av to elementer, navnet på kunden og et kontobeløp i kroner

Du skal gjøre 2 ting:
* Fullfør funksjonen `def nok_til_usd` som tar et beløp i kroner og gir tilbake beløpet i dollar, dersom man antar at 1 USD er 10.5 NOK
* Deretter skal du iterere over kundelisten og legge til kontobeløpet i USD bakerst i listen for hver kunde</text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[kundeliste = [
   ["Arve Tellefsen", 4234589],
   ["Tellef Arvesen", 23094849],
   ["Kundenavn Navnesen", 30934972978],
   ["Kari Nordmann", 239784]
]

def nok_til_usd(kr):
   kurs = 10.5
   return ...
   
###]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[kundeliste = [
   ["Arve Tellefsen", 4234589],
   ["Tellef Arvesen", 23094849],
   ["Kundenavn Navnesen", 30934972978],
   ["Kari Nordmann", 239784]
]

def nok_til_usd(kr):
   kurs = 10.5
   return kr/kurs
   
for kunde in kundeliste:
    usd_belop = nok_til_usd(kunde[1])
    kunde.append(usd_belop)]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(kundeliste)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[['Arve Tellefsen', 4234589, 403294.1904761905], ['Tellef Arvesen', 23094849, 2199509.4285714286], ['Kundenavn Navnesen', 30934972978, 2946187902.6666665], ['Kari Nordmann', 239784, 22836.571428571428]]
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 350  -->
  <question type="coderunner">
    <name>
      <text>Loops: Trekk fra 1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text>Trekk ifra 1 fra alle elementene i listen `a` og lagre de i en liste `b` ved bruk av en `for`-løkke</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text>* Vi lager først en tom liste `b = []`
* Deretter itererer vi over `a`, trekker fra 1 og legger svaret til i `b` med `b.append(tall-1)`</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload>a = [1,2,3,4,5,6,7]</answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[import re
print_tmp = print

log = []
def print_hijack(*args):
    for arg in args:
        log.append(arg)

print = print_hijack

{{ STUDENT_ANSWER }}

print = print_tmp

__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

exec_string = re.search(r'a *= *(\[.*\])', __student_answer__)
new_ans = re.sub(exec_string.group(1).replace("[", "\\[").replace("]", "\\]"), "[6,3,8,32,78,4]", __student_answer__)

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>a = [1,2,3,4,5,6,7]
b = []
for tall in a:
   b.append(tall-1)
</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print_tmp(b)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[0, 1, 2, 3, 4, 5, 6]
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>exec(new_ans)
print_tmp(b)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[5, 2, 7, 31, 77, 3]
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>